# Inverse Problems Exercises: 2022s s09 (non-sc)
https://www.umm.uni-heidelberg.de/miism/

## Notes
* Please **DO NOT** change the name of the `.ipynb` file. 
* Please **DO NOT** import extra packages to solve the tasks. 
* Please put the `.ipynb` file directly into the `.zip` archive without any intermediate folder. 

## Please provide your personal information
* full name (Name): 

<!-- grade_id: cell-2d2bf14a246580e7, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## P03: Regularization parameter estimation

```python
import numpy as np
import matplotlib.pyplot as plt
```

```python
file_gaussian = 'file_gaussian.npz'
with np.load(file_gaussian) as data:
    f_true = data['f_true']
    A_psf = data['A_psf']
    list_gn = data['list_gn']
```

### Imaging model
The imaging model can be represented by
$$
g = h \otimes f_\text{true}
= Af_\text{true}
= \mathcal{F}^{-1}\{ \mathcal{F}\{h\} \mathcal{F}\{f_\text{true}\} \},
$$
$$
g' = g + \epsilon.
$$
* $f_\text{true}$ is the input signal
* $h$ is the point spread function (kernel)
* $\otimes$ is the convolution operator
* $A$ is the Toeplitz matrix of $h$
* $\mathcal{F}$ and $\mathcal{F}^{-1}$ are the Fourier transform operator and inverse Fourier transform operator
* $\epsilon$ is the additive Gaussian noise
* $g$ is the filtered signal
* $g'$ is the noisy signal

### Mean squared error
Implement the mean squared error (MSE)
$$
\operatorname{MSE}(f)=\frac{1}{n}\sum_{i=1}^n(f_i - f_{\text{true}i})^2
$$
* Given the input signal $f$
* Given the true signal $f_\text{true}$
* Implement the function `mean_squared_error()` (using `numpy.array`)

<!-- grade_id: cell-82dbcbcd7a00edfa, points: 0, grade: false, solution: true -->
```python
def mean_squared_error(f, f_true):
    """ Compute the mean squared error comparing to the true signal:

    :param f: Input signal.
    :param f_true: True signal.
    :returns: Mean squared error.
    """
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: a001, points: 1, grade: true, solution: false -->
```python
# This cell contains hidden tests.

```

### Tikhonov regularization
Implement the objective function with Tikhonov regularization
$$
L(f) = \|Af - g'\|_2^2 + \lambda\|D'f\|_2^2,
$$
with the corresponding risidual norm $\rho = \|Af - g\|_2^2$ and solution norm $\eta = \|D'f\|_2^2$.

* Given the input signal $f$
* Given the system matrix $A$
* Given the measurement $g'$
* Given the regularization matrix $D'$
* Given the regularization parameter $\lambda$
* Return the objective function value $L(f)$ as the first output
* Return the risidual norm $\rho$ as the second output
* Return the solution norm $\eta$ as the third output
* Implement the function `objective_tikhonov()`

Implement the closed form solution of the regularized objective function
$$
\tilde f = (A^T A + \lambda D'^T D')^{-1} A^T g' = A_\lambda^{PI} g'
$$
* Given the system matrix $A$
* Given the measurement $g'$
* Given the regularization matrix $D'$
* Given the regularization parameter $\lambda$
* Implement the function `solution_tikhonov()` (using `numpy.array`)

<!-- grade_id: cell-653a2b77dc633bae, points: 0, grade: false, solution: true -->
```python
def objective_tikhonov(f, A, gn, D, lb):
    """ Compute the objective function with Tikhonov regularization.
    
    :param f: Current estimate of the signal.
    :param A: 2D matrix of the linear problem.
    :param gn: Observed signal.
    :param D: 2D matrix in the regularization term.
    :param lb: Regularization parameter.
    :returns: Objective function value, the corresponding residual norm and solution norm.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
def solution_tikhonov(A, gn, D, lb):
    """ Compute the estimate of the true signal with Tikhonov regularization.

    Use a regularization term to suppress noise.

    :param A: 2d matrix A of the linear problem.
    :param gn: Observed signal.
    :param D: 2D matrix in the regularization term.
    :param lb: Regularization parameter.
    :returns: Estimate of the true signal.
    """
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: a004, points: 3, grade: true, solution: false -->
```python
# This cell contains hidden tests.

```

<!-- grade_id: a005, points: 1, grade: true, solution: false -->
```python
# This cell contains hidden tests.

```

### Basic solutions
The minimal length solution is the solution with $D' = I$
* Calculate the closed form solution for the noisy signals in `list_gn`
* Return the outputs with $\lambda$ of $0.01$
* Save the solutions in the variable `list_f_closed` (as `list` of `numpy.array`)
* Save the corresponding objective values in the variable `list_L_closed` (as `list` of scalars)

Display the result
* Plot the outputs in `list_f_closed` in the same order of the parameter options in the subplots of `axs`
* Plot the corresponding noisy signal in each subplot
* Plot the input signal `f_true` in each subplot
* Show the legend in each subplot
* Show the case information in the titles to the subplots
* Show the mean squared error of each output comparing to `f_true` in the titles to the subplots
* Show the objective function value of each output in the titles to the subplots

<!-- grade_id: a010, points: 5, grade: true, solution: true -->
```python
fig, axs = plt.subplots(1, 3, figsize=(15, 5))
fig.suptitle('Minimal length solution')

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: a015, points: 5, grade: true, solution: false -->
```python
# This cell contains hidden tests.

```

<!-- grade_id: a020, points: 5, grade: true, solution: false -->
```python
# This cell contains hidden tests.

```

### L-curve

See: https://www.sintef.no/globalassets/project/evitameeting/2005/lcurve.pdf

The L-curve is a log-log plot of the solution norm versus the corresponding residual norm, i.e. $(\rho, \eta)$. 
The curvature of the L-curve, as a function of $\lambda$, is given by
$$
\kappa(\lambda) = 2 \frac{\hat\rho' \hat\eta'' - \hat\rho'' \hat\eta'}{((\hat\rho')^2 + (\hat\eta')^2) ^{3/2}},
$$
with
$$
\hat\rho = \log\rho,\ \ \ \ \ \hat\eta = \log\eta,
$$
where $\hat\rho'$, $\hat\eta'$, $\hat\rho''$ and $\hat\eta''$ are the first and second derivatives of $\hat\rho$ and $\hat\eta$ with respect to $\lambda$.
* Given the arrays of $\rho$, $\eta$ and $\lambda$
* Return the curvature by the definition above (using `numpy.gradient()`)
* Implement the function `calculate_curvature()`

Process the signal
* Use 100 different $\lambda \in [10^{-5}, 10]$ spaced evenly on a log scale
* Calculate the closed form solution of the noisy signal with different $\lambda$, respectively
* Save the risidual norm of the closed form solutions in the variable `list_rho` (as `list` of `numpy.array`)
* Save the solution norm of the closed form solutions in the variable `list_eta` (as `list` of `numpy.array`)
* Save the curvature of the L-curve in the variable `list_curvature` (as `list` of `numpy.array`)
* Save the mean squared error of the closed form solutions in the variable `list_mse` (as `list` of `numpy.array`)
* Save the $\lambda$ corresponding to the maximum curvature in the variable `list_lambda_l` (as `list`)

Display the result
* Show the cases of the same noisy signal in the same subplot column of `axs`
* In the first subplot row, show the scatter points of the L-curve with the residual norm in the $x$ axis and the solution norm in the $y$ axis both in log scaling
* In the first subplot row, encode the color of the scatter points by the $\lambda$ value
* In the first subplot row, show the color bar
* In the second subplot row, show the curvature of the L-curve as a function of $\lambda$ with $\lambda$ in log scaling
* In the third subplot row, show the mean squared error as a function of $\lambda$ with $\lambda$ in log scaling
* Highlight the case corresponding to the maximum curvature by a red point with a legend label in all subplots
* Add proper labels to the subplots
* Add proper titles to the subplots

<!-- grade_id: a035, points: 15, grade: true, solution: true -->
```python
def calculate_cuvature(rho_values, eta_values, lambda_values):
    """ Compute the curvature of the L-curve

    :param rho_values: Array of rho values.
    :param eta_values: Array of eta values.
    :param lambda_values: Array of lambda values.
    :returns: Array of the curvature values.
    """
# YOUR CODE HERE
raise NotImplementedError()

fig, axs = plt.subplots(3, 3, figsize=(15, 15))
fig.suptitle('L-curve')

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: a040, points: 5, grade: true, solution: false -->
```python
# This cell contains hidden tests.

lb_steps_REF = 100

np.testing.assert_array_equal(len(list_rho), len(list_gn))

for j, rho_j in enumerate(list_rho):
    print(j, end = ' ')
    np.testing.assert_array_equal(rho_j.size, lb_steps_REF)

```

<!-- grade_id: a045, points: 5, grade: true, solution: false -->
```python
# This cell contains hidden tests.

lb_steps_REF = 100

np.testing.assert_array_equal(len(list_eta), len(list_gn))

for j, eta_j in enumerate(list_eta):
    print(j, end = ' ')
    np.testing.assert_array_equal(eta_j.size, lb_steps_REF)

```

<!-- grade_id: a050, points: 5, grade: true, solution: false -->
```python
# This cell contains hidden tests.

lb_steps_REF = 100

np.testing.assert_array_equal(len(list_curvature), len(list_gn))

for j, curvature_j in enumerate(list_curvature):
    print(j, end = ' ')
    np.testing.assert_array_equal(curvature_j.size, lb_steps_REF)

```

<!-- grade_id: a055, points: 5, grade: true, solution: false -->
```python
# This cell contains hidden tests.

lb_steps_REF = 100

np.testing.assert_array_equal(len(list_mse), len(list_gn))

for j, mse_j in enumerate(list_mse):
    print(j, end = ' ')
    np.testing.assert_array_equal(mse_j.size, lb_steps_REF)

```

<!-- grade_id: a060, points: 5, grade: true, solution: false -->
```python
# This cell contains hidden tests.

np.testing.assert_array_equal(len(list_lambda_l), len(list_gn))

```

### Normalized cumulative periodogram

See https://link.springer.com/article/10.1007/s10543-006-0042-7

* Given the input signal $f$
* Given the system matrix $A$
* Given the measurement $g'$
* Calculate the normalized cumulative periodogram (NCP) according to the paper (using `numpy.fft.fft()`)
* Implement the function `compute_ncp()`

Process the signal
* Use 20 different $\lambda \in [10^{-5}, 10]$ spaced evenly on a log scale
* Calculate the closed form solution of the noisy signal with different $\lambda$, respectively
* Save the normalized cumulative periodogram (NCP) in the variable `list_ncp` (as `list` of `numpy.array`)
* Save the NCP error as the Euclidean distance between the estimated NCP and the theoretical NCP of additive Gaussian noise in the variable `list_error` (as `list` of `numpy.array`)
* Save the mean squared error of the closed form solutions in the variable `list_mse` (as `list` of `numpy.array`)
* Save the $\lambda$ corresponding to the minimum NCP error in the variable `list_lambda_l` (as `list`)

Display the result
* Show the cases of the same noisy signal in the same subplot column of `axs`
* In the first subplot row, show NCP curves
* In the first subplot row, show the legend
* In the second subplot row, show the NCP error as a function of $\lambda$ with $\lambda$ in log scaling
* Highlight the case corresponding to the minimum NCP error by a red point with a legend label in all subplots
* Add proper labels to the subplots
* Add proper titles to the subplots

<!-- grade_id: cell-06d478783038b592, points: 40, grade: true, solution: true -->
```python
def compute_ncp(f, A, gn):
    """Compute the Normalized cummulative periodigram (NCP) described by:

    P.C. Hansen et al: Exploiting residual information in the parameter choice
    for discrete ill-posed problems. Num. Math. (2006), 46: 41-59

    :param f: Current estimate of the signal.
    :param A: 2D matrix of the linear problem.
    :param gn: Observed signal.
    :return: NCP.
    """
# YOUR CODE HERE
raise NotImplementedError()

fig, axs = plt.subplots(2, 3, figsize=(15, 10))
fig.suptitle('Normalized cumulative periodogram')

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bdfc03bf5e039322, points: 5, grade: true, solution: false -->
```python
# This cell contains hidden tests.

lb_steps_REF = 20

np.testing.assert_array_equal(len(list_ncp), len(list_gn))

for j, ncp_j in enumerate(list_ncp):
    print(j, end = ' ')
    np.testing.assert_array_less(0, ncp_j.size)

```

<!-- grade_id: cell-1e431ea3239d7e77, points: 5, grade: true, solution: false -->
```python
# This cell contains hidden tests.

lb_steps_REF = 20

np.testing.assert_array_equal(len(list_error), len(list_gn))

for j, error_j in enumerate(list_error):
    print(j, end = ' ')
    np.testing.assert_array_equal(error_j.size, lb_steps_REF)

```

<!-- grade_id: cell-36a939e78827e876, points: 5, grade: true, solution: false -->
```python
# This cell contains hidden tests.

lb_steps_REF = 20

np.testing.assert_array_equal(len(list_mse), len(list_gn))

for j, mse_j in enumerate(list_mse):
    print(j, end = ' ')
    np.testing.assert_array_equal(mse_j.size, lb_steps_REF)

```

<!-- grade_id: cell-7d701ab1fa55de3f, points: 5, grade: true, solution: false -->
```python
# This cell contains hidden tests.

np.testing.assert_array_equal(len(list_lambda_ncp), len(list_gn))

```